<a href="https://colab.research.google.com/github/RomaZhm/ColabNotes/blob/main/integrals_razhmurin.ipynb" target="_parent"><img src="https://colab.research.google.com/assets/colab-badge.svg" alt="Open In Colab"/></a>

Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = "Zhmurin Roman"
COLLABORATORS = "razhmurin@edu.hse.ru"
```

---

# Вычисление определенных интегралов

## I. Двухточечная квадратура Гаусса

Постройте квадратурную формулу Гаусса, интегрирующую точно многочлены степеней вплоть до третьей на интервале $[a, b]$. Заметим, что для этого достаточно построить _двухточечную_ квадратуру.

Напишите функцию, которая возвращает веса, $w_1$ и $w_2$, и узлы, $x_1$ и $x_2$, двухточечной квадратурной формулы Гаусса с весовой функцией $w(x) = 1$, т.е., интегралы вида

$$
\int_a^b\! f(x)\, dx \approx w_1 f(x_1) + w_2 f(x_2)
$$

```python
import numpy as np

def gauss_2(a, b):
    """Return nodes and weights for a two-point Gauss quadrature on [a, b].
    
    Parameters
    ----------
    a, b : float
       Boundaries of the integration interval
       
    Returns
    -------
    x1, x2, w1, w2 : float
       Nodes and weights of the quadrature.
    """
    a1 = - 1/np.sqrt(3)
    a2 = 1/np.sqrt(3)
    w1, w2 = (b - a)/2, (b - a)/2
    x1 = (b - a)*a1 / 2 + (b + a) / 2
    x2 = (b - a)*a2 / 2 + (b + a) / 2
    return x1, x2, w1, w2
```

<!-- grade_id: cell-b6739a7f6045a9ec, points: 5, grade: true, solution: false -->
```python
from numpy.testing import assert_allclose

x1, x2, w1, w2 = gauss_2(0, 1)

def f(x, n): 
    return x**n

for n in [0, 1, 2, 3]:
    assert_allclose(w1*f(x1, n=n) + w2*f(x2, n=n),
                    1./(n+1), atol=1e-14)
```

## II. Вычитание интегрируемых сингулярностей


Вычислите определённый интеграл методом трапеций с вычитанием сингулярности

$$
I = \int_{0}^{1}\frac{e^x}{\sqrt{x(1-x)}}dx.
$$

Вам могут пригодиться значения следующих определенных интегралов:

$$
\int_0^1 \frac{1}{\sqrt{x (1-x)}} \, dx=\pi,\quad \int_0^1 \frac{x}{\sqrt{x (1-x)}} \, dx=\pi/2.
$$

Преобразуйте данный интеграл, вычитая сингулярности. Выпишите расчетные формулы.
(5 баллов, защита)

Cоставьте функцию, возвращающую значение интеграла методом трапеций.

```python
from math import sqrt, e
def integ(npts=10):
    b = 1
    a = 0
    func = (lambda x: (np.exp(x) - 1 - x) / sqrt(x*(1 - x)))
    func_0 = (lambda x: (x ** (3/2)) / 2)
    #func_1 = (lambda x: e*sqrt(1 - x)/sqrt(x) + (e * (1 - x) ** (3/2))/(2 * sqrt(x)))
    func_1 = (lambda x: e*sqrt(1-x)/2 + ((3*e - 2)*(1-x)**(3/2))/8 - ((e - 12)*(1-x)**(5/2))/48)
    h = (b - a) / npts
    I1_0 = (func_0(a + h) + func_0(a)) * h / 2
    I1_1 = (func_1(b) + func_1(b - h)) * h / 2
    I1 = I1_0
    I2 = np.pi + np.pi/2
    for j in range(2, npts):
        x_c = a + h * j
        x_p = x_c - h
        I1 += (func(x_p) + func(x_c)) * h / 2
    I = I1 + I2
    return I

```

```python
from scipy.integrate import quad
f = (lambda x: np.exp(x)/ np.sqrt(x*(1 - x)))
for n in [1000, 10000, 100000, 50000000]:
    print(n, ':', integ(npts = n), quad(f, 0, 1,))
```

Output:
```
1000 : 5.463932564208847 (5.5084297738863635, 2.04049488417013e-09)
10000 : 5.494349865751378 (5.5084297738863635, 2.04049488417013e-09)
100000 : 5.503977042180681 (5.5084297738863635, 2.04049488417013e-09)
50000000 : 5.508230640312696 (5.5084297738863635, 2.04049488417013e-09)
```

<!-- grade_id: cell-e24dd8dc92f2c8bd, points: 5, grade: true, solution: false -->
```python
# this is a test to check your computed value
from numpy.testing import assert_allclose
f = (lambda x: np.exp(x)/ np.sqrt(x*(1 - x)))
for n in [5, 10, 100, 1000]:
    assert_allclose(integ(npts = n),
                    quad(f, 0, 1)[0], atol=50/n)
```

## III. Интеграл от осциллирующей функции

Рассмотрим определенный интеграл

$$
I = \int_0^\infty\! \frac{\sin(x) \cos{(\cos{(x)})}}{x}\,dx
$$

Вычислите значение данного интеграла с относительной точностью $10^{-6}$. Для упрощения задачи можете воспользоваться функционалом `scipy.integrate.quad`.

Заметим, что "из коробки" интеграл вычислить не удается, и нужно что-то придумать.

```python
from math import sin, cos, pi

from scipy.integrate import quad
quad(lambda x: sin(x) * cos(cos(x)) / x, 0, float('inf'))
```

Output:
```
/usr/local/lib/python3.7/dist-packages/ipykernel_launcher.py:4: IntegrationWarning: The integral is probably divergent, or slowly convergent.
  after removing the cwd from sys.path.
(1.9653912540956746, 4.089174284042278)
```

Напишите функцию, которая возвращает значение данного интеграла. Только само значение интеграла, без оценки погрешности.
(оборачиваем интеграл в функцию только для удобства автопроверки).

<!-- grade_id: cell-ddd4e0b6d148fcfd, points: 0, grade: false, solution: true -->
```python
from scipy.integrate import quad
import mpmath as mp

def trap(A, p, npts=10):
    b = A
    a = 0.000001
    func = (lambda x: sin(p/x)*(x**3))
    h = (b - a) / npts
    I = 0
    for j in range(1, npts + 1):
        x_c = a + h * j
        x_p = x_c - h
        I += (func(x_p) + func(x_c)) * h / 2
    return I

def integ():
    """Return a single float, the computed value of the integral."""
    A = pi/2 + 1/654
    I1 = sin(cos(A))/A
    I2_sub = [-sin(A)/A**2 + 2*cos(A)/A**3 + 6*sin(A)/A**4 + 24 * trap(1/A, p = 1, npts = 10000),
    sin(3*A)/(3*A**2) - 2*cos(3*A)/(9*A**3) - 2*sin(3*A)/(9*A**4) - 8 * trap(1/A, p = 3, npts = 10000)/9,
    -sin(5*A)/(5*A**2) + 2*cos(5*A)/(25*A**3) + 6*sin(5*A)/(125*A**4) + 24 * trap(1/A, p = 5, npts = 10000)/125,
    sin(7*A)/(7*A**2) - 2*cos(7*A)/(49*A**3) - 6*sin(7*A)/(343*A**4) - 24 * trap(1/A, p = 7, npts = 10000)/343]
    I2 = 0
    for k in range(1, 5):
        I2 += 2 * mp.besselj((2*k - 1), 1) * I2_sub[k-1]
    I = I1 + I2
    return I

    
```

```python
print(integ())
```

Output:
```
1.18634521645079
```

<!-- grade_id: cell-b341c2ce36bc957d, points: 10, grade: true, solution: false -->
```python
from numpy.testing import assert_allclose
# this is a test cell, keep it intact
```

```python

```